# CS 3110/5110: Data Privacy
## In-Class Exercise, week of 8/26/2024

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Construct a de-identified version of the `adult` dataset.

<!-- grade_id: cell-4e7a3cdbee680759, points: 0, grade: false, solution: true -->
```python
def deidentify_adult():
    # goal: remove identifying information
    # process: 1. decide which columns have identifying info 2. remove them
    return adult.drop(['Name', 'DOB', 'SSN'], axis=1)

adult_deid = deidentify_adult()
```

<!-- grade_id: cell-68632b6467ef7923, points: 5, grade: true, solution: false -->
```python
assert 'Name' not in adult_deid.columns
assert 'DOB' not in adult_deid.columns
assert 'SSN' not in adult_deid.columns
```

## Question 2

What columns did you remove from the dataset to de-identify it? Why did you choose these columns? Why did you not choose the remaining columns?

<!-- grade_id: cell-d81a131204234a9d, points: 2, grade: true, solution: true -->
- We removed `Name`, `DOB`, and `SSN` because they are clearly identifying info
- The other columns didn't seem to contain info that would uniquely identify an individual, so we didn't remove them. But it's not clear if this is really true.

## Question 3

Write a pandas expression to return just the row containing information about Brenn McNeely.

<!-- grade_id: cell-4baa91fd73c5802e, points: 0, grade: false, solution: true -->
```python
def get_brenns_row():
    # Expression that uses the series to index into the original dataframe
    return adult[adult['Name'] == 'Brenn McNeely']

brenns_row = get_brenns_row()[['Name', 'DOB', 'Zip', 'Age']]
```

<!-- grade_id: cell-13b75b0658ee9750, points: 2, grade: true, solution: false -->
```python
assert len(brenns_row) == 1
assert brenns_row['Zip'].iloc[0] == 95668
assert brenns_row['Age'].iloc[0] == 38
```

```python
brenns_row
```

Output:
```
Name       DOB    Zip  Age
2  Brenn McNeely  8/6/1991  95668   38
```

## Question 4

Conduct a linking attack to recover Brenn's data from the `adult_deid` dataset.

<!-- grade_id: cell-99305ea08eafa471, points: 0, grade: false, solution: true -->
```python
def recover_brenns_row():
    # Goal: link brenns_row with adult_deid to find Brenn's row in adult_deid
    # brenns_row is the auxiliary data
    # adult_deid is the de-identified data
    # Idea: match rows based on zip code and age
    brenns_zip = 95668
    brenns_age = 38
    brenns_recovered_row = adult_deid[(adult_deid['Zip'] == brenns_zip) & 
                                      (adult_deid['Age'] == brenns_age)]
    return brenns_recovered_row

brenns_recovered_row = recover_brenns_row()
```

<!-- grade_id: cell-b4432ec45f4faf9c, points: 5, grade: true, solution: false -->
```python
assert len(brenns_recovered_row) == 1
assert brenns_recovered_row['Education'].iloc[0] == 'HS-grad'
```

```python
brenns_recovered_row
```

Output:
```
Zip  Age Workclass  fnlwgt Education  Education-Num Marital Status  \
2  95668   38   Private  215646   HS-grad              9       Divorced   

          Occupation   Relationship   Race   Sex  Capital Gain  Capital Loss  \
2  Handlers-cleaners  Not-in-family  White  Male             0             0   

   Hours per week        Country Target  
2              40  United-States  <=50K
```